# HW01: Các câu lệnh Linux, Python

(Cập nhật lần cuối 08/11/2025)

Họ tên: Thạch Ngọc Hân

MSSV: 23127361

---

## Cách làm bài và nộp bài

&#9889; Bạn lưu ý là mình sẽ dùng chương trình hỗ trợ chấm bài nên bạn cần phải tuân thủ chính xác qui định mà mình đặt ra, nếu không rõ thì hỏi, chứ không nên tự tiện làm theo ý của cá nhân.

**Cách làm bài**

Bạn sẽ làm trực tiếp trên file notebook này. Đầu tiên, bạn điền họ tên và MSSV vào phần đầu file ở bên trên. Trong file, bạn làm bài ở những chỗ có ghi là:
```python
# YOUR CODE HERE
raise NotImplementedError()
```
hoặc đối với những phần code không bắt buộc thì là:
```python
# YOUR CODE HERE (OPTION)
```
hoặc đối với markdown cell thì là:
```markdown
YOUR ANSWER HERE
```
Tất nhiên, khi làm thì bạn xóa dòng `raise NotImplementedError()` đi.
Đối những phần yêu cầu code thì thường ở ngay phía dưới sẽ có một (hoặc một số) cell chứa các bộ test để giúp bạn biết đã code đúng hay chưa; nếu chạy cell này không có lỗi gì thì có nghĩa là qua được các bộ test. Trong một số trường hợp, các bộ test có thể sẽ không đầy đủ; nghĩa là, nếu không qua được test thì là code sai, nhưng nếu qua được test thì chưa chắc đã đúng hoàn toàn.

Trong khi làm bài, bạn có thể cho in ra màn hình, tạo thêm các cell để test. Nhưng khi nộp bài thì bạn xóa các cell mà bạn tự tạo, xóa hoặc comment các câu lệnh in ra màn hình. Bạn lưu ý <font color=red>không được tự tiện xóa các cell hay sửa code của Thầy</font> (trừ những chỗ được phép sửa như đã nói ở trên).

Bạn cố gắng giữ cho code rõ ràng, ngăn nắp: đặt tên biến gợi ý, comment ở những chỗ cần comment, giữ cho một dòng code không quá dài (vì nếu không thì người đọc sẽ phải mất công scroll, bạn có thể dùng qui tắc: một dòng code chỉ nên có tối đa là 79 ký tự). 

Trong khi làm bài, thường xuyên `Ctrl + S` để lưu lại bài làm của bạn, tránh mất mát thông tin.


*Nên nhớ mục tiêu chính ở đây là <font color=green>học, học một cách chân thật</font>. Bạn có thể thảo luận ý tưởng với bạn khác, nhưng <font color=green>code và bài làm phải là của bạn, dựa trên sự hiểu thật sự của bạn</font>. <font color=red>Nếu vi phạm thì sẽ bị 0 điểm cho toàn bộ môn học.</font>*

**Cách nộp bài**

Khi chấm bài, đầu tiên mình sẽ chọn `Kernel` - `Restart & Run All` để restart và chạy tất cả các cell trong notebook của bạn; do đó, trước khi nộp bài, bạn nên chạy thử `Kernel` - `Restart & Run All` để đảm bảo mọi chuyện diễn ra đúng như mong đợi.

Sau đó, bạn tạo thư mục nộp bài theo cấu trúc sau:
- Thư mục `MSSV` (vd, nếu bạn có MSSV là 1234567 thì bạn đặt tên thư mục là `1234567`)
    - File `HW01.ipynb` (không cần nộp các file khác)

Cuối cùng, bạn nén thư mục `MSSV` này lại và nộp ở link trên moodle. <font color=red>Bạn lưu ý tuân thủ chính xác qui định nộp bài.</font>

---

## Môi trường code

Ta thống nhất trong môn này:
- Code trên môi trường Linux (bạn đã setup ở HW0)
- Dùng Jupyter Notebook hoặc Jupyter Lab để code file .ipynb. Lưu ý: bạn <font color=red>không dùng Google Colab</font>, vì Google Colab có thể xóa một số thông tin meta mà chương trình hỗ trợ chấm bài của mình đã tạo trong file bài tập
- Dùng phiên bản các package như trong file "min_ds-env.yml" (bạn xem trong HW0 cách dùng Conda để tạo môi trường code từ file .yml và cách kích hoạt môi trường này sau khi đã tạo)

```python
import sys
sys.executable
```

Output:
```
'/home/kimkh/HW01/.venv/bin/python'
```

Nếu không có vấn đề gì thì file chạy python sẽ là file của môi trường code "min_ds-env".

Trong bài này, ngoài Git & Jupyter Notebook/Lab thì bạn chỉ dùng Python, không dùng các thư viện khác như Numpy, Pandas, ... Trong Python, khi đọc/ghi file csv thì bạn cũng không được dùng các thư viện hỗ trợ như thư viện `csv`.

---

## Import các thư viện cần thiết

```python
import datetime
# YOUR CODE HERE (OPTION)
```

---

## Thu thập dữ liệu (3.5đ)

Trong bài tập này, bạn sẽ lấy dữ liệu log của [thùng chứa nbdime](https://github.com/namlnhut/nbdime). Nbdime là công cụ giúp Git có thể diff và merge file Notebook một cách phù hợp (thay vì chỉ coi như là một file text thông thường). Link mình cung cấp không phải là link đến thùng chứa gốc của nbdime mà là thùng chứa mình "fork" (đại khái là tạo ra một bản copy) từ thùng chứa gốc. Lý do: mình muốn cố định lại dữ liệu log để có thể cung cấp các bộ test để bạn phần nào biết là đã code đúng hay chưa (còn nếu dùng thùng chứa gốc thì dữ liệu log có thể sẽ thay đổi trong vài ngày tới, các bộ test lúc đó có thể sẽ không còn đúng nữa).

Để có thể lấy dữ liệu log thì việc đầu tiên bạn cần làm ở cell bên dưới là: 

- Dùng câu lệnh Git để clone thùng chứa về; bạn clone vào một thư mục cùng cấp với file notebook này. <font color=gray>Gợi ý: trong Jupyter Notebook, ở code cell, bạn có thể thực hiện một câu lệnh Linux (bao gồm câu lệnh Git) bằng cách thêm dấu ! ở đầu.</font>
- Sau đó, để xem thử thư mục vừa clone có nặng không, bạn tính dung lượng tổng cộng của thư mục này và lưu vào biến `repo_size` (`repo_size` là một con số, đơn vị là MB); ta thống nhất: (i) dùng cờ `--apparent-size` khi dùng câu lệnh `du`, (ii) code chỉ cần chạy đúng với thùng chứa cụ thể này. <font color=gray>Gợi ý: trong Jupyter Notebook, bạn có thể gán kết quả của câu lệnh Linux vào biến Python.</font> \
@ Các bạn dùng Mac: trong Mac thì mình nghe nói là câu lệnh `du` không có cờ `--apparent-size` (do câu lệnh `du` của Mac là bản BSD, còn của Linux là bản GNU); bạn có thể xem thử cách cài câu lệnh `du` bản GNU ở [video này](https://youtu.be/VGgTmxXp7xQ?t=894) (trong video là nói về câu lệnh `grep` nhưng mình nghĩ cũng sẽ tương tự với câu lệnh `du`).

```python
!git clone https://github.com/namlnhut/nbdime nbdime_repo
```

Output:
```
fatal: destination path 'nbdime_repo' already exists and is not an empty directory.
```

<!-- grade_id: cell-5984c6f1d3bb0311, points: 0.5, grade: true, solution: true -->
```python
# YOUR CODE HERE
tmp = !du -sh --apparent-size nbdime_repo | awk '{print $1}'
repo_size = tmp[0]
print(f"Repository size: {repo_size}")
```

Output:
```
Repository size: 21M
```

Sau khi đã clone thùng chứa về, việc kế tiếp của bạn là dùng các câu lệnh Linux (trong đó có các câu lệnh Git) và Python để lấy dữ liệu log rồi lưu xuống file "log.tsv" (định dạng tsv tương tự như csv nhưng các giá trị được cách nhau bởi tab). Để dễ hình dung, bạn có thể dùng Jupyter Lab để mở xem file kết quả đúng "correct_log.tsv" mà mình cung cấp. Như bạn có thể thấy, mỗi dòng ứng với một commit (không tính dòng đầu - dòng chứa các tên cột) và gồm các thông tin:
- "Id" là id (hash) dạng viết gọn của commit
- "ParentIds" là id dạng viết gọn của (các) commit cha
- "AuthorName" là tên của tác giả
- "AuthorEmail" là email của tác giả
- "AuthorDate" là ngày tác giả thực hiện commit, ngày được format theo chuẩn RFC2822 (xem phần gợi ý ở bên dưới để biết cách lấy ngày theo format này) 
- "Subject" là tựa đề của message (message tương tự như email, ngoài phần "subject" thì có thể có thêm phần "body"; ở đây ta chỉ lấy phần "subject")
- "ChangedFiles": các file bị thay đổi, nếu có nhiều file thì các file được phân tách bởi dấu phẩy

Tại sao mình dùng định dạng tsv thay vì csv? Vì dấu phẩy nhiều khả năng là sẽ có trong phần "subject" của message. Tab ít khả năng sẽ có hơn, nhưng nếu có thì bạn hãy thay bằng bốn khoảng trắng.

<font color=gray>Gợi ý cho phần này:</font> 
- <font color=gray>Bạn có thể xem thêm về câu lệnh `git log` [ở đây](https://git-scm.com/book/en/v2/Git-Basics-Viewing-the-Commit-History)</font>
- <font color=gray>Khi ở Notebook, bạn thử `cd` đến thư mục clone rồi `pwd` xem? Kết quả như vậy là vì khi thực hiện một dòng lệnh Linux ở Notebook thì Notebook sẽ tạo ra một terminal tạm thời để thực thi dòng lệnh đó trong phút chốc. Như vậy, câu lệnh `cd` vẫn là có `cd` đến thư mục clone nhưng là ở terminal tạm thời; khi câu lệnh `cd` chạy xong thì terminal tạm thời này sẽ được đóng và ta vẫn ở thư mục mà hiện chứa file Notebook. Để có thể `cd` rồi `pwd` thì một cách là ghi 2 câu lệnh này trên một dòng và cách nhau bởi `&&`; lúc này thì cả 2 câu lệnh sẽ được thực thi trên cùng một terminal</font>

<!-- grade_id: cell-7a31931ee1dd0b8b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
git_log_output = !cd nbdime_repo && git log --pretty=format:"---COMMIT_START---%n%h%n%p%n%an%n%ae%n%aD%n%s" --name-only
log_data_str = "\n".join(git_log_output)
commit_blocks = log_data_str.split("---COMMIT_START---")[1:]

output_filename = "log.tsv"
header = "Id\tParentIds\tAuthorName\tAuthorEmail\tAuthorDate\tSubject\tChangedFiles\n"

with open(output_filename, "w", encoding="utf-8") as f:
    f.write(header)
    
    for block in commit_blocks:
        block = block.strip() 
        if not block:
            continue 
        
        lines = block.split('\n')
        if len(lines) < 6:
            continue 
        
        commit_id = lines[0].strip()
        parent_ids = lines[1].strip()
        author_name = lines[2].strip()
        author_email = lines[3].strip()
        author_date = lines[4].strip()
        subject = lines[5].strip()
        
        subject = subject.replace('\t', '    ')
        file_list = [line.strip() for line in lines[6:] if line.strip()]
        files_str = ",".join(file_list)
        
        f.write(f"{commit_id}\t{parent_ids}\t{author_name}\t{author_email}\t{author_date}\t{subject}\t{files_str}\n")
```

<!-- grade_id: cell-1671c3f106d5da7d, points: 3, grade: true, solution: false -->
```python
# TEST
file = open('log.tsv', 'r')
log = file.read()
file.close()
file = open('correct_log.tsv', 'r')
correct_log = file.read()
file.close()
assert log == correct_log
```

---

## Khám phá dữ liệu

### Dữ liệu có bao nhiêu dòng và bao nhiêu cột? (0.5đ)

Đầu tiên, bạn sẽ đọc dữ liệu từ file "log.tsv" (nếu phần thu thập dữ liệu bạn làm chưa đúng thì bạn có thể đọc từ file "correct_log.tsv" mà mình đã cung cấp). Bạn lưu kết quả đọc được vào biến `cols`. Biến `cols` là một dictionary, trong đó mỗi phần tử ứng với một cột, key là chuỗi ứng với tên cột, value là list chứa các giá trị của cột (mỗi giá trị bạn cứ tạm để kiểu dữ liệu là chuỗi).

Kế tiếp, từ biến `cols`, bạn sẽ tính số dòng và số cột, rồi lần lượt lưu vào hai biến là `num_rows` và `num_cols`.

<!-- grade_id: cell-f2f6dbf283e0ccee, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
filename = "log.tsv"
with open(filename, 'r', encoding='utf-8') as f:
    content = f.read()

lines = content.strip().split('\n')
header = lines[0].split('\t')
data_lines = lines[1:]

cols = {col_name: [] for col_name in header}

for line in data_lines:
    values = line.split('\t')
    for i, col_name in enumerate(header):
        if i < len(values):
            cols[col_name].append(values[i])
        else:
            cols[col_name].append("") 

num_rows = len(data_lines)
num_cols = len(header)
```

```python
# TEST
assert num_rows == 2149
assert num_cols == 7
assert set(cols.keys()) == set(['Id', 'ParentIds', 'AuthorName', 
                                'AuthorEmail', 'AuthorDate', 'Subject', 
                                'ChangedFiles'])
id_col_len = len(cols['Id'])
for col_name in cols:
    assert len(cols[col_name]) == id_col_len
```

### Mỗi dòng có ý nghĩa gì? Có vấn đề các dòng có ý nghĩa khác nhau không?

Mỗi dòng ứng với một commit. Có vẻ không có dòng nào "lạc loài" (nghĩa là không có vấn đề các dòng có ý nghĩa khác nhau).

### Dữ liệu có các dòng bị lặp không? (0.25đ)

Có lẽ là không có. Để chắc ăn thì một cách là ta có thể check cột "Id"; nếu cột "Id" không có giá trị bị lặp thì có thể kết luận là không có dòng bị lặp. Bạn sẽ làm vụ này và lưu kết quả check vào biến `duplicated_id`; biến này bằng True nếu cột "Id" có giá trị lặp và bằng False nếu ngược lại.

<!-- grade_id: cell-31e053e964e6c9eb, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
id_list = cols['Id']
unique_ids = set(id_list)
duplicated_id = len(unique_ids) != len(id_list)
```

<!-- grade_id: cell-ab5957518b489fa7, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert duplicated_id == False
```

### Mỗi cột có ý nghĩa gì?

Vụ này mình đã nói lúc thu thập dữ liệu. Ở đây, mình ghi lại một lần nữa:

- "Id" là id (hash) dạng viết gọn của commit
- "ParentIds" là id dạng viết gọn của (các) commit cha
- "AuthorName" là tên của tác giả
- "AuthorEmail" là email của tác giả
- "AuthorDate" là ngày tác giả thực hiện commit
- "Subject" là tựa đề của message (message tương tự như email, ngoài phần "subject" thì có thể có thêm phần "body"; ở đây ta chỉ lấy phần "subject")
- "ChangedFiles": các file bị thay đổi, nếu có nhiều file thì các file được phân tách bởi dấu phẩy

### Mỗi cột hiện đang có kiểu dữ liệu gì? Có cột nào có kiểu dữ liệu chưa phù hợp để có thể xử lý tiếp không?

- Các cột "Id", "ParentIds", "AuthorName", "AuthorEmail", "Subject": kiểu dữ liệu dạng chuỗi là đúng rồi
- Cột "ChangedFiles": hmm ... có lẽ tạm thời ta vẫn để dạng chuỗi, có gì cần thì ta sẽ xử lý sau
- Cột "AuthorDate": cần chuyển từ kiểu dữ liệu dạng chuỗi sang datetime; để có thể tiếp tục khám phá tiếp thì ngay sau đây bạn sẽ tiền xử lý vụ này luôn

---

## Tiền xử lý (0.5đ)

Ở đây, bạn sẽ chuyển cột "AuthorDate" từ dạng chuỗi sang datetime. <font color=gray>Gợi ý: bạn xem [document](https://docs.python.org/3/library/datetime.html#strftime-and-strptime-format-codes) để biết cách truyền format phù hợp cho hàm `strptime`.</font>

<!-- grade_id: cell-742e74182a9d1936, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
date_format = "%a, %d %b %Y %H:%M:%S %z"

converted_dates = []
for date_str in cols['AuthorDate']:
    dt = datetime.datetime.strptime(date_str.strip(), date_format)
    converted_dates.append(dt)

cols['AuthorDate'] = converted_dates
```

<!-- grade_id: cell-c551b6a15162bb10, points: 0.5, grade: true, solution: false -->
```python
# TEST
assert type(cols['AuthorDate'][0]) == datetime.datetime
assert cols['AuthorDate'][0] == datetime.datetime(2025, 10, 9, 11, 32, 46, 
            tzinfo=datetime.timezone(datetime.timedelta(seconds=7200)))
```

---

## Khám phá dữ liệu (tiếp tục)

### Với mỗi cột có kiểu dữ liệu dạng số (numeric) và dạng datetime, các giá trị được phân bố như thế nào? (0.5đ)

Ở đây thì không có cột dạng số, chỉ có một cột dạng datetime là cột "AuthorDate".

Đầu tiên, bạn tính tỉ lệ % các giá trị thiếu của cột "AuthorDate". <font color=gray>Gợi ý: không cần code, có thể lập luận dựa vào kết quả khi chuyển cột "AuthorDate" từ dạng chuỗi sang datetime ở bên trên.</font>

<!-- grade_id: cell-47e2584daf9dda74, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE

Không có giá trị nào bị lỗi khi chuyển sang `datetime`
→ **Tỉ lệ giá trị thiếu: 0%**

Tiếp theo, bạn tính min và max của cột "AuthorDate", rồi lần lượt lưu vào hai biến `author_date_min` và `author_date_max`.

<!-- grade_id: cell-ed5647763d70e921, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
author_date_min = min(cols["AuthorDate"])
author_date_max = max(cols["AuthorDate"])
```

<!-- grade_id: cell-8bb55737a75c0a6e, points: 0.25, grade: true, solution: false -->
```python
# TEST
assert author_date_min == datetime.datetime(2015, 11, 9, 13, 13, 28, 
            tzinfo=datetime.timezone(datetime.timedelta(seconds=3600)))
assert author_date_max == datetime.datetime(2025, 10, 9, 11, 32, 46, 
            tzinfo=datetime.timezone(datetime.timedelta(seconds=7200)))
```

Hmm ... Min và max của cột "AuthorDate" không có gì bất thường.

### Với mỗi cột có kiểu dữ liệu dạng phân loại (categorical), các giá trị được phân bố như thế nào? (1.75đ)

Các cột dạng phân loại là các cột mà không phải là cột "AuthorDate". Bạn sẽ lưu thông tin của các cột này vào biến `cate_col_profiles`. Biến này là một dictionary, trong đó mỗi phần tử ứng với một cột dạng phân loại, key là tên cột, value là list gồm: tỉ lệ % giá trị thiếu, số lượng các giá trị khác nhau (không tính giá trị thiếu), set chứa các giá trị khác nhau (không bao gồm giá trị thiếu).

<!-- grade_id: cell-d9eb895996d7b157, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
categorical_cols = [col for col in cols if col != "AuthorDate"]
cate_col_profiles = {}

for col in categorical_cols:
    values = cols[col]
    num_missing = sum(1 for v in values if v == "")
    missing_ratio = num_missing / len(values) * 100
    unique_values = set(v for v in values if v != "")
    num_unique = len(unique_values)
    cate_col_profiles[col] = [missing_ratio, num_unique, unique_values]
```

<!-- grade_id: cell-50123f76f79e1f27, points: 1, grade: true, solution: false -->
```python
# TEST
assert set(cate_col_profiles.keys()) == {'AuthorEmail', 'AuthorName', 
                                         'ChangedFiles', 'Id', 
                                         'ParentIds', 'Subject'}
for col_name, col_profile in cate_col_profiles.items():
    assert len(col_profile) == 3
    assert len(col_profile[2]) == col_profile[1]
assert round(cate_col_profiles['ParentIds'][0], 3) == 0.047
assert round(cate_col_profiles['ChangedFiles'][0], 3) == 15.914
assert cate_col_profiles['ParentIds'][1] == 2029
assert cate_col_profiles['ChangedFiles'][1] == 934
```

```python
# TEST (BY EYES)
print(f"{'ColName':12} {'Miss(%)':7} {'NumDifVals':10} {'SomeVals'}")
for col_name, col_profile in cate_col_profiles.items():
    print(f'{col_name:12} {col_profile[0]:<7.3f} {col_profile[1]:<10} {col_profile[2].__repr__()[:34]+"...":<}')
```

Output:
```
ColName      Miss(%) NumDifVals SomeVals
Id           0.000   2149       {'2cd6efe', 'fd69229', '8ca8c6f', ...
ParentIds    0.047   2029       {'fd69229', '8ca8c6f', 'fba5c31', ...
AuthorName   0.000   55         {'Tony Hirst', 'Todd', 'Nicholas B...
AuthorEmail  0.000   58         {'piotr.chromiec@rtbhouse.com', 't...
Subject      0.000   2058       {'Merge pull request #372 from min...
ChangedFiles 15.914  934        {'nbdime/tests/test_js_artifacts_i...
```

Bạn quan sát kết quả in ra ở trên (tạm thời chưa để ý đến cột "Miss(%)") và cho biết có gì bất thường không? Hay mọi thứ đều bình thường? Cho mình biết tại sao bạn lại kết luận như vậy.

<!-- grade_id: cell-88055a6ed4991805, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE

Cột Id hầu như tất cả giá trị đều duy nhất → bình thường, đúng đặc tính của ID.

Cột ParentIds cũng gần như duy nhất, chỉ có vài giá trị trùng lặp → bình thường.

Cột AuthorName và AuthorEmail có số lượng giá trị hạn chế (55 và 58) → bình thường cho danh sách tác giả.

Cột Subject có gần như mỗi commit một giá trị khác nhau (2058/2149) → bất thường nếu muốn phân nhóm, nhưng hợp lý vì mỗi commit thường có subject riêng.

Kết luận:

Về tổng thể, dữ liệu categorical không có lỗi lớn, nhưng cột Subject hầu như không thể dùng để nhóm vì mỗi giá trị gần như duy nhất.

Nguyên nhân: đặc tính dữ liệu commit, mỗi subject là riêng biệt.

Kế tiếp, bạn hãy cho in ra các dòng mà có giá trị thiếu (in sao là tùy bạn, miễn là dễ nhìn). Sau đó, bạn hãy quan sát các dòng này để xem tại sao lại bị thiếu (trong trường hợp, bạn không biết tại sao thì cứ ghi nhận là không biết tại sao).

<!-- grade_id: cell-fda56cd71a2bb407, points: 0.25, grade: true, solution: true -->
```python
# YOUR CODE HERE
missing_rows_idx = [
    i for i in range(len(cols["Id"])) 
    if any(cols[col][i] == "" for col in cols)
]

for i in missing_rows_idx:
    row = {col: cols[col][i] for col in cols}
    print(f"Row {i+1}: {row}")
```

Output:
```
Row 1: {'Id': '606b6fd', 'ParentIds': 'c2d2cc2 55a4d01', 'AuthorName': 'M Bussonnier', 'AuthorEmail': 'bussonniermatthias@gmail.com', 'AuthorDate': datetime.datetime(2025, 10, 9, 11, 32, 46, tzinfo=datetime.timezone(datetime.timedelta(seconds=7200))), 'Subject': 'Merge pull request #789 from Carreau/master-main', 'ChangedFiles': ''}
Row 5: {'Id': 'ecf8996', 'ParentIds': '030509a 5e10953', 'AuthorName': 'M Bussonnier', 'AuthorEmail': 'bussonniermatthias@gmail.com', 'AuthorDate': datetime.datetime(2025, 10, 8, 13, 39, 18, tzinfo=datetime.timezone(datetime.timedelta(seconds=7200))), 'Subject': 'Merge pull request #767 from jupyter/dependabot/github_actions/dot-github/workflows/actions/download-artifact-4.1.7', 'ChangedFiles': ''}
Row 25: {'Id': '3dcdce3', 'ParentIds': '2f54b80 8806423', 'AuthorName': 'Vidar Tonaas Fauske', 'AuthorEmail': 'vidartf@gmail.com', 'AuthorDate': datetime.datetime(2023, 11, 10, 19, 18, 44, tzinfo=datetime.timezone.utc), 'Subject': 'Merge pull request #734 from ju
```

<!-- grade_id: cell-0ce04cd6c1dd7a31, points: 0.25, grade: true, solution: true -->
YOUR ANSWER HERE

Tất cả các dòng có giá trị thiếu đều nằm ở cột ChangedFiles.

Nguyên nhân hợp lý: những dòng này là các commit merge (thường là merge pull request), nên không có thay đổi trực tiếp trong file, dẫn đến ChangedFiles trống.

Không có giá trị thiếu ở các cột quan trọng khác như Id, AuthorName, AuthorEmail, AuthorDate.

### Kết luận:

Việc thiếu dữ liệu ở ChangedFiles là bình thường do đặc tính commit merge.

Không phát hiện nguyên nhân bất thường khác.

---

## Đưa ra câu hỏi có ý nghĩa cần trả lời

Sau khi khám phá dữ liệu, ta đã hiểu hơn về dữ liệu. Bây giờ, ta hãy xem thử có câu hỏi nào có ý nghĩa mà có thể được trả lời bằng dữ liệu này không.

Các câu hỏi có thể có là: 
- *Câu hỏi 1: trong năm 2021, ai là người có nhiều commit nhất, ai là người có nhiều commit nhì, ...?* \
Tìm câu trả lời cho câu hỏi này để làm gì? \
Hmm ... Chẳng hạn để khen thưởng
- *Câu hỏi 2: commit nào là commit cuối cùng (xét trên "AuthorDate") đã thay đổi file "nbdime/webapp/templates/difftool.html"?* \
Tìm câu trả lời cho câu hỏi này để làm gì? \
Hmm ... Chẳng hạn để đổ tội cho tác giả của commit đó ;-)

---

## Tiền xử lý (nếu cần) + phân tích dữ liệu để trả lời cho câu hỏi 1 (1đ)

Nhắc lại câu hỏi 1: *trong năm 2023, ai là người có nhiều commit nhất, ai là người có nhiều commit nhì, ...?*

Bạn sẽ lưu kết quả vào biến `wanted_people`. Biến này là một dictionary, mỗi phần tử ứng với một người mà có commit trong năm 2023, key là tên của người đó, value là số lượng commit của người đó. Dictionary này được sort theo value giảm dần.

<!-- grade_id: cell-a40eed2ac02f6c2b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
commit_2023_idx = [i for i in range(len(cols["AuthorDate"]))
                   if cols["AuthorDate"][i].year == 2023]

author_counts = {}
for i in commit_2023_idx:
    author = cols["AuthorName"][i]
    if author in author_counts:
        author_counts[author] += 1
    else:
        author_counts[author] = 1
wanted_people = dict(sorted(author_counts.items(), key=lambda x: x[1], reverse=True))
```

<!-- grade_id: cell-9e312d16b92a0e40, points: 1, grade: true, solution: false -->
```python
# TEST
correct_wanted_people = {'Vidar Tonaas Fauske': 62,
                         'Frédéric Collonval': 31,
                         'Florence Haudin': 18,
                         'github-actions[bot]': 14,
                         'fcollonval': 4,
                         'dependabot[bot]': 4,
                         'HaudinFlorence': 4,
                         'krassowski': 3,
                         'Michał Krassowski': 3,
                         'Chih-Hsuan Yen': 2,
                         'Angus Hollands': 1,
                         'Min RK': 1,
                         'Johan Mabille': 1,
                         'Saiwing Yeung': 1,
                         'Cory Gwin': 1}
assert wanted_people == correct_wanted_people
```

---

## Tiền xử lý (nếu cần) + phân tích dữ liệu để trả lời cho câu hỏi 2 (2đ)

Nhắc lại câu hỏi 2: *commit nào là commit cuối cùng (xét trên "AuthorDate") đã thay đổi file "nbdime/webapp/templates/difftool.html"?*

Ta thống nhất: không xét các commit "merge pull request" vì các commit này không có thông tin về "ChangedFiles".

Bạn lưu kết quả vào biến `wanted_commit`. Biến này là một dictionary, trong đó mỗi phần tử ứng với một thông tin của commit, key là tên cột, value là giá trị tương ứng.

<!-- grade_id: cell-37a6a90f597ad4a1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
target_file = "nbdime/webapp/templates/difftool.html"
latest_date = None
latest_commit_index = -1

num_commits = len(cols['Id'])
for i in range(num_commits):
    files_str = cols['ChangedFiles'][i]
    if files_str == "":
        continue
    
    individual_files = files_str.split(',')
    if target_file in individual_files:
        current_date = cols['AuthorDate'][i]
        if latest_date is None or current_date > latest_date:
            latest_date = current_date
            latest_commit_index = i

wanted_commit = {}
if latest_commit_index != -1:
    for col_name in cols:
        wanted_commit[col_name] = cols[col_name][latest_commit_index]
```

<!-- grade_id: cell-41c49732faf571f7, points: 2, grade: true, solution: false -->
```python
# TEST
correct_wanted_commit = {'Id': 'd39f45a',
 'ParentIds': '8ebbc6e',
 'AuthorName': 'Frédéric Collonval',
 'AuthorEmail': 'fcollonval@users.noreply.github.com',
 'AuthorDate': datetime.datetime(2023, 10, 20, 23, 20, 31, tzinfo=datetime.timezone(datetime.timedelta(seconds=7200))),
 'Subject': 'Do not collapse editor lines in single editor (#719)',
 'ChangedFiles': 'nbdime/webapp/templates/compare.html,nbdime/webapp/templates/diff.html,nbdime/webapp/templates/difftool.html,nbdime/webapp/templates/merge.html,nbdime/webapp/templates/mergetool.html,packages/nbdime/src/common/mergeview.ts,ui-tests/data/merge_test6/center.ipynb,ui-tests/data/merge_test6/left.ipynb,ui-tests/data/merge_test6/right.ipynb,ui-tests/package-lock.json,ui-tests/tests/nbdime-merge-test6.spec.ts,ui-tests/tests/nbdime-merge-test6.spec.ts-snapshots/merge-test6-take-a-snapshot-at-opening-1-linux.png,ui-tests/tests/nbdime-merge-test6.spec.ts-snapshots/merge-test6-take-a-snapshot-at-opening-1-win32.png'}
assert wanted_commit == correct_wanted_commit
```